# Assignment 1

Nested Data and Nested Iteration

The variable `nested` contains a nested list. Assign `"snake"` to the variable `output` **using indexing**. Do not hard code.

<!-- grade_id: cell-bdd016730873b172, points: 0, grade: false, solution: true -->
```python
nested = [['dog', 'cat', 'horse'], ['frog', 'turtle', 'snake', 'gecko'], ['hamster', 'gerbil', 'rat', 'ferret']]

output = None # Replace "None" here with an expression that indexes into `nested`


output = nested[1][2]  # indexing to get "snake"


output
```

Output:
```
'snake'
```

<!-- grade_id: cell-c830fb33be763368, points: 1, grade: true, solution: false -->
```python
# make the following assert pass; using indexing, not hardcoding.
assert output == 'snake'

```

A nested list (`lst`) is provided below. Use indexing and the `in` operator to create variables with Boolean values. See comments for further instructions. Do not hard code.

<!-- grade_id: cell-f846a83774f43929, points: 0, grade: false, solution: true -->
```python
lst = [['apple', 'orange', 'banana'], [5, 6, 7, 8, 9.9, 10], ['green', 'yellow', 'purple', 'red']]

# Use indexing and the 'in' operator to assign values to the variables according to the instructions below
# No hardcoding!

# Test to see if 'yellow' is in the third list of lst. Save to variable `yellow`
yellow = None # Replace "None" here with a boolean expression

# Test to see if 4 is in the second list of lst. Save to variable `four`
four = None # Replace "None" here with a boolean expression

# Test to see if 'orange' is in the first element of lst. Save to variable `orange`
orange = None # Replace "None" here with a boolean expression



lst = [['apple', 'orange', 'banana'], [5, 6, 7, 8, 9.9, 10], ['green', 'yellow', 'purple', 'red']]

yellow = 'yellow' in lst[2]  # Check if 'yellow' is in the 3rd sublist
four = 4 in lst[1]           # Check if 4 is in the 2nd sublist
orange = 'orange' in lst[0]  # Check if 'orange' is in the 1st sublist

print(yellow, four, orange)
```

Output:
```
True False True
```

<!-- grade_id: cell-612bd76129dd8a88, points: 1, grade: true, solution: false -->
```python
assert yellow == True, "yellow does not have the correct value"
assert four == False, "four does not have the correct value"
assert orange == True, "orange does not have the correct value"

```

A nested list (`L`) is provided below. Use indexing and the `in` operator to create variables with Boolean values. See comments for further instructions. Do not hard code.

<!-- grade_id: cell-6f00defc22a9a644, points: 0, grade: false, solution: true -->
```python
L = [[5, 8, 7], ['hello', 'hi', 'hola'], [6.6, 1.54, 3.99], ['small', 'large']]

# Check if 'hola' is in the list L. Save to variable name test1
test1 = None # Replace "None" here with a boolean expression

# Check if [5, 8, 7] is in the list L. Save to variable name test2
test2 = None # Replace "None" here with a boolean expression

# Check if 6.6 is in the third element of list L. Save to variable name test3
test3 = None # Replace "None" here with a boolean expression


# Check if 'hola' is in the second sublist of L (i.e., L[1])
test1 = 'hola' in L

# Check if the first sublist [5, 8, 7] is in L
test2 = L[0] in L

# Check if 6.6 is in the third sublist (i.e., L[2])
test3 = 6.6 in L[2]

print(test1, test2, test3)


```

Output:
```
False True True
```

<!-- grade_id: cell-e808d6b4735e5efa, points: 1, grade: true, solution: false -->
```python
assert test1 == False, "test1 does not have the correct value"
assert test2 == True,  "test2 does not have the correct value"
assert test3 == True,  "test3 does not have the correct value"

```

A nested list (`nested`) is provided below. Use indexing and the `in` operator to create variables with Boolean values. See comments for further instructions. Do not hard code.

<!-- grade_id: cell-c4c23e947c4fc73a, points: 0, grade: false, solution: true -->
```python
nested = {
    'data': ['finding', 23, ['exercises', 'hangout', 34]],
    'window': ['part', 'whole', [], 'sum', ['math', 'calculus', 'algebra', 'geometry', 'statistics', ['physics', 'chemistry', 'biology']]]
}

# Check to see if the string 'data' is a key in nested. If it is, assign True to the variable data, otherwise assign False.
data = None # Replace "None" here with a boolean expression

# Check to see if the integer 24 is in the value of the key 'data' in nested. If it is, assign True to the variable twentyfour, otherwise assign False.
twentyfour = None # Replace "None" here with a boolean expression

# Check to see that the string 'whole' is *not* in the value of the key 'window' in nested. If it's not, then assign True to the variable whole, otherwise assign False.
whole = None # Replace "None" here with a boolean expression

# Check to see if the string 'physics' is a key in the dictionary nested. If it is, assign True to the variable physics, otherwise False.
physics = None # Replace "None" here with a boolean expression


# Check if 'data' is a key in the dictionary
data = 'data' in nested

# Check if 24 is in the list assigned to the key 'data'
twentyfour = 24 in nested['data']

# Check if 'whole' is NOT in the list assigned to 'window'
whole = 'whole' not in nested['window']

# Check if 'physics' is a key in the dictionary
physics = 'physics' in nested


print(data, twentyfour, whole, physics)
```

Output:
```
True False False False
```

<!-- grade_id: cell-0da054b06b2edd59, points: 1, grade: true, solution: false -->
```python
assert data == True, "data does not have the correct value"
assert twentyfour == False, "twentyfour does not have the correct value"
assert whole == False, "whole does not have the correct value"
assert physics == False, "physics does not have the correct value"

```

The variable `nested_d` contains a nested dictionary with the gold medal counts for the top four countries in the past three Olympics. Assign the value of Great Britain’s gold medal count from the London Olympics to the variable `london_gold`. Use indexing. Do not hardcode.

<!-- grade_id: cell-4d810eef76aead78, points: 0, grade: false, solution: true -->
```python
nested_d = { 
    'Beijing': {'China':51, 'USA':36, 'Russia':22, 'Great Britain':19},
    'London':  {'USA':46, 'China':38, 'Great Britain':29, 'Russia':22},
    'Rio':     {'USA':35, 'Great Britain':22, 'China':20, 'Germany':13}
}

london_gold = None # Replace "None" here with an expression that indexes into `nested_d`


# Use nested dictionary indexing
london_gold = nested_d['London']['Great Britain']

london_gold
```

Output:
```
29
```

<!-- grade_id: cell-8590b00f64d313ec, points: 1, grade: true, solution: false -->
```python
assert london_gold == 29

```

Below, we have provided a nested dictionary, `sports`. Index into the dictionary to create variables that we have listed in the code window.

<!-- grade_id: cell-f69ccbbacd2993e4, points: 0, grade: false, solution: true -->
```python
sports = {
    'swimming': ['butterfly', 'breaststroke', 'backstroke', 'freestyle'],
    'diving':   ['springboard', 'platform', 'synchronized'],
    'track':    ['sprint', 'distance', 'jumps', 'throws'],
    'gymnastics': {
        'women':['vault', 'floor', 'uneven bars', 'balance beam'],
        'men':  ['vault', 'parallel bars', 'floor', 'rings']
    }
}

# Assign the string 'backstroke' to the variable v1
v1 = None # Replace "None" here with an expression that indexes into `sports`

# Assign the string 'platform' to the variable v2
v2 = None # Replace "None" here with an expression that indexes into `sports`

# Assign the list ['vault', 'floor', 'uneven bars', 'balance beam'] to the variable v3
v3 = None # Replace "None" here with an expression that indexes into `sports`

# Assign the string 'rings' to the variable v4
v4 = None # Replace "None" here with an expression that indexes into `sports`



# Assign the string 'backstroke' to the variable v1
v1 = sports['swimming'][2]

# Assign the string 'platform' to the variable v2
v2 = sports['diving'][1]

# Assign the list ['vault', 'floor', 'uneven bars', 'balance beam'] to the variable v3
v3 = sports['gymnastics']['women']

# Assign the string 'rings' to the variable v4
v4 = sports['gymnastics']['men'][3]


print(v1, v2, v3, v4)
```

Output:
```
backstroke platform ['vault', 'floor', 'uneven bars', 'balance beam'] rings
```

<!-- grade_id: cell-a0971634a44eb834, points: 1, grade: true, solution: false -->
```python
assert v1 == 'backstroke'
assert v2 == 'platform'
assert v3 == ['vault', 'floor', 'uneven bars', 'balance beam']
assert v4 == 'rings'

```

Create a new list called `thirds` by iterating through the contents of `list__of_lists`, collecting the third element of each sublist.

<!-- grade_id: cell-ddda847974ac8019, points: 0, grade: false, solution: true -->
```python
list__of_lists = [
    ['purple', 'mauve', 'blue'],
    ['red', 'maroon', 'blood orange', 'crimson'],
    ['sea green', 'cornflower', 'lavender', 'indigo'],
    ['yellow', 'amarillo', 'mac n cheese', 'golden rod']
]

thirds = None



thirds = [sublist[2] for sublist in list__of_lists]

thirds
```

Output:
```
['blue', 'blood orange', 'lavender', 'mac n cheese']
```

<!-- grade_id: cell-1dd02ab20b37a271, points: 1, grade: true, solution: false -->
```python
assert thirds == ['blue', 'blood orange', 'lavender', 'mac n cheese']

```

Given below is a list of lists of athletes. Create a list, `t_athletes`, that includes an athlete's name if it contains a lowercase "t". If it does not contain the letter "t", save that athlete's name into list `others`.

<!-- grade_id: cell-f14dc189e95de246, points: 0, grade: false, solution: true -->
```python
athletes = [
    ['Phelps', 'Lochte', 'Schooling', 'Ledecky', 'Franklin'],
    ['Felix', 'Bolt', 'Gardner', 'Eaton'],
    ['Biles', 'Douglas', 'Hamm', 'Raisman', 'Mikulak', 'Dalton']]

t_athletes = None
others = None


t_athletes = []
others = []

for sublist in athletes:
    for athlete in sublist:
        if 't' in athlete:
            t_athletes.append(athlete)
        else:
            others.append(athlete)

```

<!-- grade_id: cell-be34ffa738a0764b, points: 1, grade: true, solution: false -->
```python

```